```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02B.ipynb")
```

# Lab 02B: Data Cleaning and EDA

To receive credit for a lab, answer all questions correctly and submit before the deadline.

You must submit this assignment to Pensieve by the on-time deadline, **Tuesday, February 10th, 11:59 PM PT**. As a reminder, slip days are **not** applicable on labs. **We strongly encourage you to plan to submit your work to Pensieve several hours before the stated deadline.** This way, you will have ample time to contact staff for submission support.  

## Lab Walk-Through
In addition to the lab notebook, we have also released a prerecorded walk-through video of the lab. We encourage you to reference this video as you work through the lab. Run the cell below to display the videos.


**Note:** This video is recorded in Fall 2023. There may be slight inconsistencies between the version you are viewing versus the version used in the recording, but content is identical.

[**Link for the lab walk-through**](https://www.youtube.com/watch?v=hwRYs5ZRgW4&list=PLQCcNQgUcDfqSg049DVFZCQbupMY5Bn5Z)

```python
from IPython.display import YouTubeVideo
YouTubeVideo("hwRYs5ZRgW4", list = 'PLQCcNQgUcDfqSg049DVFZCQbupMY5Bn5Z', listType = 'playlist')
```

## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others, please **include their names** below. (That's a good way to learn your classmates' names.)

**Collaborators**: *list collaborators here*

## Discussion 3 Mini-lecture

In Data 100, discussions will not reserve time to host mini-lectures. Instead, we will release a set of pre-recorded mini-lectures that supplement the concepts introduced in lecture. The week after a mini-lecture is released, the topics will be covered in discussion. Thus, it is important to watch these mini-lectures **before** attending the discussion section you are assigned to. Discussion 3 mini-lecture is attached below. [**Link for the Discussion 3 Mini-lecture**](https://www.youtube.com/watch?v=9jFqjbPLThc&list=PLQCcNQgUcDfq8wQjfrvDVmaLa2jl2R1WK&index=3).


**Note**: This video is recorded in Fall 2024, now quantitative continuous, quantitative discrete, and keys are no longer in scope.


<span style="color: red; font-size: 150%;">**You must be logged into your Berkeley account on YouTube to access the video**</span>


```python
YouTubeVideo("9jFqjbPLThc", list = 'PLQCcNQgUcDfq8wQjfrvDVmaLa2jl2R1WK', listType='playlist')
```

---
### Debugging Guide
If you run into any technical issues, we highly recommend checking out the [Data 100 Debugging Guide](https://ds100.org/debugging-guide/). In this guide, you can find general questions about Jupyter notebooks / Datahub, Gradescope, and common pandas errors.

---
[`pandas`](https://pandas.pydata.org/) is one of the most widely used `Python` libraries in data science. In this lab, you will review commonly used data-wrangling operations/tools in `pandas`. We continue the content from the previous lab and aim to give you familiarity with:

* Aggregating the data (using `.groupby`),
* Filtering the data (using boolean arrays and `groupby.filter`),
* Pivoting (using `.pivot_table`).

In this lab, you are going to use several `pandas` methods. Reminder from the lecture that you may press `shift+tab` on method parameters to see the documentation for that method. For example, if you were using the `drop` method in `pandas`, you could press `shift+tab` to see what `drop` is expecting.

`pandas` is very similar to the `datascience` library that you saw in Data 8. This [conversion notebook](https://data100.datahub.berkeley.edu/hub/user-redirect/git-pull?repo=https%3A%2F%2Fgithub.com%2FDS-100%2Fsu23-materials&branch=main&urlpath=lab%2Ftree%2Fsu23-materials%2Flec%2Flec02%2Fdata8_translation_examples.ipynb) may serve as a useful guide!

This lab expects that you have watched all three `pandas` lectures. If you have not, this lab will probably take a very long time.

```python
import numpy as np
import pandas as pd
%matplotlib inline
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

### **REVIEW:** `Groupby` and `Groupby` Shorthand

Let's now turn to use `groupby` from lectures 3 and 4.

### Elections

Let's start by reading in the election dataset from the `pandas` lectures.

```python
# Run this cell to load data from CSV file; no further action is needed.
elections = pd.read_csv("data/elections.csv")
elections.head(5)
```

Output:
```
Year          Candidate                  Party  Popular vote Result  \
0  1824     Andrew Jackson  Democratic-Republican        151271   loss   
1  1824  John Quincy Adams  Democratic-Republican        113142    win   
2  1828     Andrew Jackson             Democratic        642806    win   
3  1828  John Quincy Adams    National Republican        500897   loss   
4  1832     Andrew Jackson             Democratic        702735    win   

           %  
0  57.210122  
1  42.789878  
2  56.203927  
3  43.796073  
4  54.574789
```

As we saw before, we can `groupby` a specific column, e.g., `"Party"` and can print out the resulting sub-DataFrames. The output below can help you get an understanding of what `groupby` is actually doing.

An example is given below for elections since 1980.

```python
# Run this cell to print sub-DataFrames of a groupby object; no further action is needed.
for n, g in elections[elections["Year"] >= 1980].groupby("Party"):
    print(f"Name: {n}") # By the way, this is an "f string", a great feature of Python
    display(g)
```

Output:
```
Name: Citizens
     Year       Candidate     Party  Popular vote Result         %
127  1980  Barry Commoner  Citizens        233052   loss  0.270182
Name: Constitution
     Year         Candidate         Party  Popular vote Result         %
160  2004  Michael Peroutka  Constitution        143630   loss  0.117542
164  2008     Chuck Baldwin  Constitution        199750   loss  0.152398
172  2016    Darrell Castle  Constitution        203091   loss  0.149640
Name: Democratic
     Year        Candidate       Party  Popular vote Result          %
129  1980     Jimmy Carter  Democratic      35480115   loss  41.132848
134  1984   Walter Mondale  Democratic      37577352   loss  40.729429
137  1988  Michael Dukakis  Democratic      41809074   loss  45.770691
140  1992     Bill Clinton  Democratic      44909806    win  43.118485
144  1996     Bill Clinton  Democratic      47400125    win  49.296938
151  2000          Al Gore  Democratic      50999897   loss  48.491813
158  2004       John Kerry
```

Recall that once we've formed groups, we can aggregate each sub-DataFrame (a.k.a. group) into a single row using an aggregation function. For example, if we use `.agg('mean')` on the groups above, we get back a single `DataFrame` where each group has been replaced by a single row. In each column for that aggregate row, the value that appears is the average of all values in that group.

For columns that are non-numeric, e.g., `"Result"`, the `pandas` version we're using (version 2.0.2) will error because we cannot compute the mean of the `Result` column. To remedy this, we add a `numeric_only=True` argument to our function calls so that we only calculate the `mean` for columns that contain numeric values. Alternatively, we can manually select only the numerical columns before calling the `agg` function so the aggregation is only applied to numerical columns.

```python
elections_after_1980 = elections[elections["Year"] >= 1980]

elections_after_1980.groupby("Party").agg('mean', numeric_only=True)

# alternatively, we can manually select only the numerical columns before calling `agg`
# elections_after_1980.groupby("Party")[['Year', 'Popular vote', '%']].agg('mean')
```

Output:
```
Year  Popular vote          %
Party                                                  
Citizens           1980.000000  2.330520e+05   0.270182
Constitution       2009.333333  1.821570e+05   0.139860
Democratic         2002.000000  5.623007e+07  47.441567
Green              2010.000000  8.803689e+05   0.740615
Independent        2004.000000  4.692752e+06   4.593141
Libertarian        2000.000000  1.026652e+06   0.830876
Libertarian Party  2024.000000  6.501300e+05   0.418895
Natural Law        1996.000000  1.136700e+05   0.118219
New Alliance       1988.000000  2.172210e+05   0.237804
Populist           1992.000000  1.061520e+05   0.101918
Reform             1998.000000  4.267094e+06   4.417831
Republican         2002.000000  5.611919e+07  48.057815
Taxpayers          1996.000000  1.846560e+05   0.192045
```

Equivalently we can use one of the shorthand aggregation functions, e.g. `.mean()`: 

```python
elections_after_1980.groupby("Party").mean(numeric_only=True)
```

Output:
```
Year  Popular vote          %
Party                                                  
Citizens           1980.000000  2.330520e+05   0.270182
Constitution       2009.333333  1.821570e+05   0.139860
Democratic         2002.000000  5.623007e+07  47.441567
Green              2010.000000  8.803689e+05   0.740615
Independent        2004.000000  4.692752e+06   4.593141
Libertarian        2000.000000  1.026652e+06   0.830876
Libertarian Party  2024.000000  6.501300e+05   0.418895
Natural Law        1996.000000  1.136700e+05   0.118219
New Alliance       1988.000000  2.172210e+05   0.237804
Populist           1992.000000  1.061520e+05   0.101918
Reform             1998.000000  4.267094e+06   4.417831
Republican         2002.000000  5.611919e+07  48.057815
Taxpayers          1996.000000  1.846560e+05   0.192045
```

Note that the index of the `DataFrame` returned by a `groupby.agg` call is no longer a set of numeric indices from $0$ to $N-1$. Instead, we see that the index for the example above is now the `Party`. If we want to restore our `DataFrame` so that `Party` is a column rather than the index, we can use `reset_index`.

```python
elections_after_1980.groupby("Party").mean(numeric_only=True).reset_index()
```

Output:
```
Party         Year  Popular vote          %
0            Citizens  1980.000000  2.330520e+05   0.270182
1        Constitution  2009.333333  1.821570e+05   0.139860
2          Democratic  2002.000000  5.623007e+07  47.441567
3               Green  2010.000000  8.803689e+05   0.740615
4         Independent  2004.000000  4.692752e+06   4.593141
5         Libertarian  2000.000000  1.026652e+06   0.830876
6   Libertarian Party  2024.000000  6.501300e+05   0.418895
7         Natural Law  1996.000000  1.136700e+05   0.118219
8        New Alliance  1988.000000  2.172210e+05   0.237804
9            Populist  1992.000000  1.061520e+05   0.101918
10             Reform  1998.000000  4.267094e+06   4.417831
11         Republican  2002.000000  5.611919e+07  48.057815
12          Taxpayers  1996.000000  1.846560e+05   0.192045
```

**IMPORTANT NOTE:** Notice that the code above consists of chained method calls. This sort of code is very common in `pandas` programming and in data science in general. Such chained method calls can sometimes go many layers deep, so you might consider adding newlines between lines of code for clarity. For example, we could instead write the code above as:

```python
# pandas method chaining
(
elections.query("Year >= 1980").groupby("Party") 
                               .mean(numeric_only=True)  ## Computes the mean values by party
                               .reset_index()            ## Resets to a numerical index
)
```

Output:
```
Party         Year  Popular vote          %
0            Citizens  1980.000000  2.330520e+05   0.270182
1        Constitution  2009.333333  1.821570e+05   0.139860
2          Democratic  2002.000000  5.623007e+07  47.441567
3               Green  2010.000000  8.803689e+05   0.740615
4         Independent  2004.000000  4.692752e+06   4.593141
5         Libertarian  2000.000000  1.026652e+06   0.830876
6   Libertarian Party  2024.000000  6.501300e+05   0.418895
7         Natural Law  1996.000000  1.136700e+05   0.118219
8        New Alliance  1988.000000  2.172210e+05   0.237804
9            Populist  1992.000000  1.061520e+05   0.101918
10             Reform  1998.000000  4.267094e+06   4.417831
11         Republican  2002.000000  5.611919e+07  48.057815
12          Taxpayers  1996.000000  1.846560e+05   0.192045
```

Note that we have surrounded the entire call by a big set of parentheses so that `Python` doesn't complain about the indentation. An alternative is to use the \ symbol to indicate to `Python` that your code continues on to the next line!

```python
# pandas method chaining (alternative)
elections[elections["Year"] >= 1980].groupby("Party") \
                               .mean(numeric_only=True) \
                               .reset_index()     
```

Output:
```
Party         Year  Popular vote          %
0            Citizens  1980.000000  2.330520e+05   0.270182
1        Constitution  2009.333333  1.821570e+05   0.139860
2          Democratic  2002.000000  5.623007e+07  47.441567
3               Green  2010.000000  8.803689e+05   0.740615
4         Independent  2004.000000  4.692752e+06   4.593141
5         Libertarian  2000.000000  1.026652e+06   0.830876
6   Libertarian Party  2024.000000  6.501300e+05   0.418895
7         Natural Law  1996.000000  1.136700e+05   0.118219
8        New Alliance  1988.000000  2.172210e+05   0.237804
9            Populist  1992.000000  1.061520e+05   0.101918
10             Reform  1998.000000  4.267094e+06   4.417831
11         Republican  2002.000000  5.611919e+07  48.057815
12          Taxpayers  1996.000000  1.846560e+05   0.192045
```

**IMPORTANT NOTE:** You should NEVER solve problems like the one above using loops or list comprehensions. This is slow and also misses the entire point of this part of Data 100. 

Before we continue, we'll print out the election dataset again for your convenience. 

```python
elections.head(5)
```

Output:
```
Year          Candidate                  Party  Popular vote Result  \
0  1824     Andrew Jackson  Democratic-Republican        151271   loss   
1  1824  John Quincy Adams  Democratic-Republican        113142    win   
2  1828     Andrew Jackson             Democratic        642806    win   
3  1828  John Quincy Adams    National Republican        500897   loss   
4  1832     Andrew Jackson             Democratic        702735    win   

           %  
0  57.210122  
1  42.789878  
2  56.203927  
3  43.796073  
4  54.574789
```

<br>

---

### Question 1a
Using `groupby.agg` or one of the shorthand methods (`groupby.min`, `groupby.first`, etc.), create a `Series` object `best_result_percentage_only` that returns a `Series` showing the entire best result for every party, sorted in decreasing order. Your `Series` should include only parties that have earned at least 10% of the vote in some election. Your result should look like this:

<code>
Party
Democratic               61.344703
Republican               60.907806
Democratic-Republican    57.210122
National Union           54.951512
Whig                     53.051213
Liberal Republican       44.071406
National Republican      43.796073
Northern Democratic      29.522311
Progressive              27.457433
American                 21.554001
Independent              18.956298
Southern Democratic      18.138998
American Independent     13.571218
Constitutional Union     12.639283
Free Soil                10.138474
Name: %, dtype: float64
</code>
<br/>

[A list of named `groupby.agg` shorthand methods](https://pandas.pydata.org/pandas-docs/stable/user_guide/groupby.html#aggregation) (you'll have to scroll down about one page).


```python
best_result_percentage_only = (
    elections.groupby("Party")
             .max(numeric_only=True)
             .sort_values("%", ascending=False)
             .pipe(lambda s: s[s > 10])
             .dropna()['%']
)
best_result_percentage_only
```

Output:
```
Party
Democratic               61.344703
Republican               60.907806
Democratic-Republican    57.210122
National Union           54.951512
Whig                     53.051213
Liberal Republican       44.071406
National Republican      43.796073
Northern Democratic      29.522311
Progressive              27.457433
American                 21.554001
Independent              18.956298
Southern Democratic      18.138998
American Independent     13.571218
Constitutional Union     12.639283
Free Soil                10.138474
Name: %, dtype: float64
```

```python
grader.check("q1a")
```

<br>

---

### Question 1b  
Repeat Question 1a. However, this time, your result should be a `DataFrame` showing all available information (all columns) rather than only the percentage as a `Series`.

This question is trickier than Question 1a. Make sure to check the `pandas` lecture slides if you're stuck! It's very easy to make a subtle mistake that shows Woodrow Wilson and William Taft both winning the 2020 election.

For example, the first 3 rows of your table should be:

|Party | Year | Candidate      | Popular Vote | Result | %         |
|------|------|----------------|--------------|--------|-----------|
|**Democratic**  | 1964 | Lyndon Johnson | 43127041      | win   | 61.344703 |
|**Republican**  | 1972 | Richard Nixon | 47168710      | win   | 60.907806 |
|**Democratic-Republican**  | 1824 | Andrew Jackson | 151271      | loss   | 57.210122 |

Note that the index is `Party`. In other words, don't use `reset_index`.


```python
best_result = (
    elections.loc[elections.groupby("Party")["%"].idxmax()]
             .sort_values("%", ascending=False)
)
best_result = best_result[best_result["%"] > 10].set_index("Party")
best_result
```

Output:
```
Year               Candidate  Popular vote Result  \
Party                                                                      
Democratic             1964          Lyndon Johnson      43127041    win   
Republican             1972           Richard Nixon      47168710    win   
Democratic-Republican  1824          Andrew Jackson        151271   loss   
National Union         1864         Abraham Lincoln       2211317    win   
Whig                   1840  William Henry Harrison       1275583    win   
Liberal Republican     1872          Horace Greeley       2834761   loss   
National Republican    1828       John Quincy Adams        500897   loss   
Northern Democratic    1860      Stephen A. Douglas       1380202   loss   
Progressive            1912      Theodore Roosevelt       4122721   loss   
American               1856        Millard Fillmore        873053   loss   
Independent            1992              Ross Perot      19743821   loss   
Southern Democratic    1860    John
```

```python
grader.check("q1b")
```

### **REVIEW:** `DataFrameGroupBy.filter`

Our `DataFrame` contains a number of parties that have never had a successful presidential run. For example, the 2020 elections included candidates from the Libertarian and Green parties, neither of which have elected a president.

```python
# Run this cell to print the last four rows; no further action is needed.
elections.tail(4)
```

Output:
```
Year       Candidate              Party  Popular vote Result          %
183  2024   Kamala Harris         Democratic      75019230   loss  48.336772
184  2024      Jill Stein              Green        861155   loss   0.554864
185  2024  Robert Kennedy        Independent        756383   loss   0.487357
186  2024    Chase Oliver  Libertarian Party        650130   loss   0.418895
```

Suppose we were conducting an analysis trying to focus our attention on parties that had elected a president. 

The most natural approach is to use `groupby.filter` [(docs)](https://pandas.pydata.org/docs/reference/api/pandas.core.groupby.DataFrameGroupBy.filter.html). This is an incredibly powerful but subtle tool for filtering data.

The code below accomplishes the task at hand. It does this by creating a function that returns `True` if and only if a sub-`DataFrame` (a.k.a. group) contains at least one winner. This function, in turn, uses the `pandas` function `any` [(docs)](https://pandas.pydata.org/docs/reference/api/pandas.Series.any.html).

```python
# Run this cell to keep only the rows of parties that have 
# elected a president; no further action is needed.
def at_least_one_candidate_in_the_frame_has_won(frame):
    """Returns df with rows only kept for parties that have
    won at least one election
    """
    return (frame["Result"] == 'win').any()

winners_only = (
    elections
        .groupby("Party")
        .filter(at_least_one_candidate_in_the_frame_has_won)
)
winners_only.tail(5)
```

Output:
```
Year        Candidate       Party  Popular vote Result          %
176  2016  Hillary Clinton  Democratic      65853514   loss  48.521539
178  2020     Joseph Biden  Democratic      81268924    win  51.311515
179  2020     Donald Trump  Republican      74216154   loss  46.858542
182  2024     Donald Trump  Republican      77303568    win  49.808629
183  2024    Kamala Harris  Democratic      75019230   loss  48.336772
```

Alternately, we could have used a `lambda` function instead of explicitly defining a named function using `def`. 

```python
# Run this cell to keep only the rows of parties that have 
# elected a president; no further action is needed.
winners_only = (
    elections
        .groupby("Party")
        .filter(lambda x : (x["Result"] == "win").any())
)
winners_only.tail(5)
```

Output:
```
Year        Candidate       Party  Popular vote Result          %
176  2016  Hillary Clinton  Democratic      65853514   loss  48.521539
178  2020     Joseph Biden  Democratic      81268924    win  51.311515
179  2020     Donald Trump  Republican      74216154   loss  46.858542
182  2024     Donald Trump  Republican      77303568    win  49.808629
183  2024    Kamala Harris  Democratic      75019230   loss  48.336772
```

<br>

---

### Question 1c

Using `filter`, create a `DataFrame` object `major_party_results_since_1988` that includes all election results starting after 1988 (exclusive) but only shows a row if the Party it belongs to has earned at least 1% of the popular vote in ANY election since 1988.

For example, despite having candidates in 2004, 2008, and 2016, no Constitution party candidates should be included since this party has not earned 1% of the vote in any election since 1988. However, you should include the Reform candidate from 2000 (Pat Buchanan) despite only having 0.43% of the vote, because in 1996, the Reform candidate Ross Perot exceeded 1% of the vote.

For example, the first three rows of the table you generate should look like:

|     |   Year | Candidate         | Party       |   Popular vote | Result   |         % |
|----:|-------:|:------------------|:------------|---------------:|:---------|----------:|
| 139 |   1992 | Andre Marrou      | Libertarian |       290087   | loss     | 0.278516  |
| 140 |   1992 | Bill Clinton      | Democratic  |       44909806 | win      | 43.118485 |
| 142 |   1992 | George H. W. Bush | Republican  |       39104550 | loss     |  37.544784|

*Hint*: The following questions might help you construct your solution. One of the lines should be identical to the `filter` examples shown above.

1) How can we **only** keep rows in the data starting after 1988 (exclusive)?
2) What column should we `groupby` to filter out parties that have earned at least 1% of the popular vote in ANY election since 1988?
3) How can we write an aggregation function that takes a sub-DataFrame and returns whether at least 1% of the vote has been earned in that sub-DataFrame? This may give you a hint about the second question!


```python
major_party_results_since_1988 = elections[elections['Year'] > 1988].groupby('Party').filter(lambda x: (x['%'] >= 1).any())
major_party_results_since_1988.head()
```

Output:
```
Year          Candidate        Party  Popular vote Result          %
139  1992       Andre Marrou  Libertarian        290087   loss   0.278516
140  1992       Bill Clinton   Democratic      44909806    win  43.118485
142  1992  George H. W. Bush   Republican      39104550   loss  37.544784
143  1992         Ross Perot  Independent      19743821   loss  18.956298
144  1996       Bill Clinton   Democratic      47400125    win  49.296938
```

```python
grader.check("q1c")
```

### **REVIEW:** `str`

`pandas` provides special purpose functions for working with specific common data types such as strings and dates, which you will learn about in more detail in Lecture 6. For example, the code below provides the length of every Candidate's name from our `elections` dataset. 

```python
elections["Candidate"].str.len()
```

Output:
```
0      14
1      17
2      14
3      17
4      14
       ..
182    12
183    13
184    10
185    14
186    12
Name: Candidate, Length: 187, dtype: int64
```

<br>

---

### Question 2

Using `.str.split`, create a new `DataFrame` called `elections_with_first_name` with a new column `First Name` that is equal to the Candidate's first name.

See the `pandas` `str` [documentation](https://pandas.pydata.org/docs/reference/api/pandas.Series.str.split.html) for documentation on using `str.split`.

```python
elections_with_first_name = elections.copy()
elections_with_first_name['First Name'] = elections_with_first_name['Candidate'].str.split().str[0]
elections_with_first_name
```

Output:
```
Year          Candidate                  Party  Popular vote Result  \
0    1824     Andrew Jackson  Democratic-Republican        151271   loss   
1    1824  John Quincy Adams  Democratic-Republican        113142    win   
2    1828     Andrew Jackson             Democratic        642806    win   
3    1828  John Quincy Adams    National Republican        500897   loss   
4    1832     Andrew Jackson             Democratic        702735    win   
..    ...                ...                    ...           ...    ...   
182  2024       Donald Trump             Republican      77303568    win   
183  2024      Kamala Harris             Democratic      75019230   loss   
184  2024         Jill Stein                  Green        861155   loss   
185  2024     Robert Kennedy            Independent        756383   loss   
186  2024       Chase Oliver      Libertarian Party        650130   loss   

             % First Name  
0    57.210122     Andrew  
1    42.789878       John  
2    56.
```

```python
grader.check("q2")
```

<br>

---

## Babynames
Remember the `babynames` dataset from Lab02A? Let's load it in again and explore the data with our newly covered functions! Like last time, we'll only load in data from California. 

Run the following cell: 

```python
file_path = 'data/namesbystate_ca.txt.gz'
column_labels = ['State', 'Sex', 'Year', 'Name', 'Count']

babynames = pd.read_csv(file_path, names=column_labels)

babynames.head()
```

Output:
```
State Sex  Year      Name  Count
0    CA   F  1910      Mary    295
1    CA   F  1910     Helen    239
2    CA   F  1910   Dorothy    220
3    CA   F  1910  Margaret    163
4    CA   F  1910   Frances    134
```

The code below creates a table with the frequency of all names from 2022. 

```python
# Run this cell to create a table with the frequency 
# of all names from 2022; no further action is needed.
babynames_2022 = (
    babynames[babynames['Year'] == 2022]
              .groupby("Name")
              .sum()[["Count"]]
              .reset_index()
)
babynames_2022
```

Output:
```
Name  Count
0     Aadhini      6
1     Aadhira      5
2      Aadhya     33
3        Aadi     11
4       Aadit      5
...       ...    ...
6196     Zyla     20
6197    Zylah     14
6198     Zylo      5
6199     Zyon     17
6200     Zyra     16

[6201 rows x 2 columns]
```

<br>

---

### Question 3

Using the `pd.merge` [(docs)](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.merge.html) function described in the lecture, combine the `babynames_2022` table with the `elections_with_first_name` table you created earlier to form `presidential_candidates_and_name_popularity`. Your resulting `DataFrame` should contain all columns from both of the tables.

```python
presidential_candidates_and_name_popularity = pd.merge(elections_with_first_name, babynames_2022, left_on="First Name", right_on="Name")
presidential_candidates_and_name_popularity
```

Output:
```
Year          Candidate                  Party  Popular vote Result  \
0    1824     Andrew Jackson  Democratic-Republican        151271   loss   
1    1824  John Quincy Adams  Democratic-Republican        113142    win   
2    1828     Andrew Jackson             Democratic        642806    win   
3    1828  John Quincy Adams    National Republican        500897   loss   
4    1832     Andrew Jackson             Democratic        702735    win   
..    ...                ...                    ...           ...    ...   
142  2020       Donald Trump             Republican      74216154   loss   
143  2020     Howard Hawkins                  Green        405035   loss   
144  2024       Donald Trump             Republican      77303568    win   
145  2024     Robert Kennedy            Independent        756383   loss   
146  2024       Chase Oliver      Libertarian Party        650130   loss   

             % First Name    Name  Count  
0    57.210122     Andrew  Andrew    741  
1    4
```

```python
grader.check("q3")
```

### **REVIEW:** `pandas.pivot_table`

Suppose we want to build a table showing the total number of babies born of each sex in each year. One way is to `groupby` using both columns of interest:

```python
babynames.groupby(["Year", "Sex"])[["Count"]].agg('sum').head(6)
```

Output:
```
Count
Year Sex       
1910 F     5950
     M     3213
1911 F     6602
     M     3381
1912 F     9804
     M     8142
```

While this does give us the information we're looking for, a more natural approach is to use pivot tables to represent our data in a more readable format.

```python
babynames_pivot = babynames.pivot_table(
    index = "Year",     # rows (turned into index)
    columns = "Sex",    # column values
    values = ["Count"], # field(s) to process in each group
    aggfunc = 'sum',   # group operation
)
babynames_pivot.head(6)
```

Output:
```
Count       
Sex       F      M
Year              
1910   5950   3213
1911   6602   3381
1912   9804   8142
1913  11860  10234
1914  13815  13111
1915  18643  17192
```

We can also include multiple values in our pivot tables 

```python
babynames_pivot = babynames.pivot_table(
    index = "Year",     # rows (turned into index)
    columns = "Sex",    # column values
    values = ["Count", "Name"],
    aggfunc = 'max',   # group operation
)
babynames_pivot.head(6)
```

Output:
```
Count          Name         
Sex      F     M       F        M
Year                             
1910   295   237  Yvonne  William
1911   390   214   Zelma   Willis
1912   534   501  Yvonne  Woodrow
1913   584   614   Zelma   Yoshio
1914   773   769   Zelma   Yoshio
1915   998  1033    Zita    Yukio
```

<br>

---

### Question 4

Using `presidential_candidates_and_name_popularity`, create a table, `party_result_popular_vote_pivot`, whose index is the `Party` and whose columns are their `Result`. Each cell should contain the total number of popular votes received.  [`pandas` 
`pivot_table` documentation](https://pandas.pydata.org/docs/reference/api/pandas.pivot_table.html).

You may notice that there are `NaN`s in your table from missing data. Replace the `NaN` values with 0. You may find `.pivot_table`'s `fill_value=` argument helpful. Or, you can use [`pd.DataFrame.fillna` (documentation here)](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.fillna.html).

```python
party_result_popular_vote_pivot = presidential_candidates_and_name_popularity.pivot_table(
    index="Party", 
    columns="Result",
    values="Count",
    aggfunc="sum", 
    fill_value=0
)
party_result_popular_vote_pivot
```

Output:
```
Result                 loss   win
Party                            
American                524     0
American Independent    733     0
Anti-Masonic            831     0
Anti-Monopoly          1524     0
Citizens                  9     0
Communist               831     0
Constitution           1041     0
Constitutional Union    490     0
Democratic             7346  5532
Democratic-Republican   741   490
Free Soil               683     0
Green                  1117     0
Greenback              1094     0
Independent            1458     0
Libertarian            2451     0
Libertarian Party       209     0
National Democratic     490     0
National Republican    1361     0
National Union            0   353
Natural Law             490     0
New Alliance             19     0
Northern Democratic      87     0
Populist               1657     0
Progressive            2595     0
Prohibition            3408     0
Reform                    7     0
Republican             6374  8016
Socialist     
```

```python
grader.check("q4")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Lab 02B!

### Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly, weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://docs.google.com/forms/d/e/1FAIpQLSelI7jXeKh6eYqWnpEsJVPSuRL0JWZ2yjLNbZJMxekUquDetA/viewform?usp=dialog). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Submit this file to the Lab 02B assignment on Pensieve. If you run into any issues when running this cell, feel free to check this [section](https://ds100.org/debugging-guide/autograder-gradescope/#why-does-the-last-grader-export-cell-fail-if-all-previous-tests-passed) in the Data 100 Debugging Guide.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: None)
```
>>> len(best_result_percentage_only)
15
```
```
>>> float(best_result_percentage_only['Independent'].sum())
18.95629754
```
```
>>> float(best_result_percentage_only.iloc[0])
61.34470329
```


Tests for q1b (points: None)
```
>>> best_result.shape
(15, 5)
```
```
>>> float(best_result['Popular vote'].sum())
135020916.0
```
```
>>> best_result.loc['Democratic', 'Candidate']
'Lyndon Johnson'
```


Tests for q1c (points: None)
```
>>> major_party_results_since_1988.shape
(41, 6)
```
```
>>> float(major_party_results_since_1988['%'].min())
0.098088334
```
```
>>> major_party_results_since_1988['Candidate'].iloc[0]
'Andre Marrou'
```


Tests for q2 (points: None)
```
>>> elections_with_first_name.shape
(187, 7)
```
```
>>> elections_with_first_name[elections_with_first_name['Candidate'] == 'Andrew Jackson'].iloc[0]['First Name']
'Andrew'
```
```
>>> elections_with_first_name[elections_with_first_name['Candidate'] == 'Jo Jorgensen'].iloc[0]['First Name']
'Jo'
```


Tests for q3 (points: None)
```
>>> presidential_candidates_and_name_popularity.shape
(147, 9)
```
```
>>> presidential_candidates_and_name_popularity[presidential_candidates_and_name_popularity['Candidate'] == 'Andrew Jackson'].iloc[0]['Name']
'Andrew'
```
```
>>> int(presidential_candidates_and_name_popularity[presidential_candidates_and_name_popularity['Candidate'] == 'William Lemke'].iloc[0]['Count'])
831
```


Tests for q4 (points: None)
```
>>> party_result_popular_vote_pivot.index.name
'Party'
```
```
>>> 'Result' in party_result_popular_vote_pivot.columns.names
True
```
```
>>> int(party_result_popular_vote_pivot.isnull().sum().sum())
0
```
```
>>> int(party_result_popular_vote_pivot.loc['American', 'Popular vote']['loss'] if isinstance(party_result_popular_vote_pivot.columns, pd.MultiIndex) else party_result_popular_vote_pivot.loc['American']['loss'])
158271
```
```
>>> int(party_result_popular_vote_pivot.loc['American', 'Popular vote']['win'] if isinstance(party_result_popular_vote_pivot.columns, pd.MultiIndex) else party_result_popular_vote_pivot.loc['American']['win'])
0
```
```
>>> int(party_result_popular_vote_pivot.loc['Democratic', 'Popular vote']['loss'] if isinstance(party_result_popular_vote_pivot.columns, pd.MultiIndex) else party_result_popular_vote_pivot.loc['Democratic']['loss'])
314190254
```
```
>>> int(party_result_popular_vote_pivot.loc['Democratic', 'Popular vote']['win'] if isinstance(party_result_popular_vote_pivot.columns, pd.MultiIndex) else party_result_popular_vote_pivot.loc['Democratic']['win'])
441751531
```
